# Assignment 1 - Sam Forman

## 1. (1.0 point) Function return value

This exercise is to show you how function return values work. Run the code below to see what happens.

```python
def f1(x):
    print(x + 1)

def f2(x):
    return(x + 1)

f1(4)
f2(4)
```

Output:
```
5
5
```

Each function should produce the same output. Now run the following code.

```python
f1(4) + 1

f2(4) + 1
```

Output:
```
5
TypeError                                 Traceback (most recent call last)
<ipython-input-11-aee8bde09b65> in <module>()
----> 1 f1(4) + 1
      2 
      3 f2(4) + 1

TypeError: unsupported operand type(s) for +: 'NoneType' and 'int'
```

Explain what is happening and why you get this result.

This error is happening because the output of f1 is a printed statement sent to the consol, which can't be combined with a interger. 

## 2. (1.0 point) Function operation

Consider the following piece of code:

```python
def f(x):
    if x == 0 or x == 1:
        return x
    return f(x - 1) + f(x - 2)
    
```

```python

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-14-611fdd6a369e> in <module>()
----> 1 print(f(1))

NameError: name 'f' is not defined
```

```python
def f(x):
    if x == 0 or x == 1:
        return x
    return f(x - 1) + f(x - 2)

print(f(1))
print(f(2))
print(f(3))
print(f(4))
print(f(5))
print(f(10))
print(f(15))
```

Output:
```
1
1
2
3
5
55
610
```

Describe, in words, what this code does, and how it does it.

This function first evalues the "x" that is passed to it on whether it is equal to 0 or 1. If it is then the original x is returned. If it is not then the function returns the sum of the whole function run on x-1 and x-2. After that point I'm not totally sure how the function terminates for larger numbers, I assume it branches out in a certain way but can't totally think through it. 

For what inputs will this function not behave as expected? What will happen?

For larger numbers (~ >5) different behavior starts happen where the function returns larger and larger numbers. Not totally sure why that happens. 

## 3. (1.0 point) Defining a function

The following program plays the game of "Rock, Paper, and Scissors". Familiarize yourself with how it works. Try running the program.

```python
from random import randint
valid_input = ["R", "P", "S"]
winners = [["R", "S"], ["S", "P"], ["P", "R"]]

print("You are playing a Rock, Paper, Scissors simulation.")
user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
user_choice = user_choice.upper()

while user_choice != 'Q':
    if user_choice not in valid_input:
        print("Error: invalid input -- " + user_choice)
        break
    
    computer_choice = valid_input[randint(0, len(valid_input)-1)]
    print("Computer chooses " + computer_choice)

    if user_choice == computer_choice:
        print("Tie Game")
    elif [user_choice, computer_choice] in winners:
        print("You win!")
    else:
        print("Computer wins!")
    
    user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
    user_choice = user_choice.upper()
```

Recall how we define a function using `def`, and how we pass in parameters. Transform the code into a function called `rps` that takes two parameters: `user_choice` and `computer_choice`, with valid inputs `'R'`, `'P'`, and `'S'`. Your function should return `'D'` for a draw, `'W'` if `user_choice` wins, and `'L'` if `user_choice` loses.

```python
### BEGIN SOLUTION

def rps(user_choice, computer_choice):
    rps_valid_input = ["R", "P", "S"]
    rps_winners = [["R", "S"], ["S", "P"], ["P", "R"]]
    
    if (user_choice or computer_choice) not in rps_valid_input:
        print("Error: invalid input -- " + user_choice)
        return 0 
    
    # first assure that input is valid
    
    #then replicate last part of the game where user_choice, comp_choice is compared to the list of winning values
        #and if it matches return win, lose, draw accordingly. 
    
    if user_choice == computer_choice:
        return("D")
    elif [user_choice, computer_choice] in rps_winners:
        return("W")
    else:
        return("L")


### END SOLUTION
```

<!-- grade_id: Q3C, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(rps('R', 'R'), 'D')
assert_equal(rps('R', 'P'), 'L')
assert_equal(rps('R', 'S'), 'W')
assert_equal(rps('P', 'R'), 'W')
assert_equal(rps('P', 'P'), 'D')
assert_equal(rps('P', 'S'), 'L')
assert_equal(rps('S', 'R'), 'L')
assert_equal(rps('S', 'P'), 'W')
assert_equal(rps('S', 'S'), 'D')
```

## 4. (1.0 point) List function

Given a list of non-empty tuples, return a list sorted in increasing order by the last element in each tuple. 

For example: `[(1, 7), (1, 3), (3, 4, 5), (2, 2)]` yields `[(2, 2), (1, 3), (3, 4, 5), (1, 7)]`

<!-- grade_id: Q4A, points: 0, grade: false, solution: true -->
```python
test_tuples = [(1, 7), (1, 3), (3, 4, 5), (2, 2)]

def sort_tuple_list(tuple_list):
    '''Function to sort a list of non-empty tuples in increasing order by the last element in each tuple'''
    ### BEGIN SOLUTION
    
    return(sorted(tuple_list, key= lambda tup : tup[-1]))
    
    
    ## key sources: https://wiki.python.org/moin/HowTo/Sorting
    ##              https://docs.python.org/3/tutorial/controlflow.html
    
    ### END SOLUTION
    

print(sort_tuple_list(test_tuples))
```

Output:
```
[(2, 2), (1, 3), (3, 4, 5), (1, 7)]
```

<!-- grade_id: Q4B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(sort_tuple_list([(1, 3), (3, 2), (2, 1)]), [(2, 1), (3, 2), (1, 3)])
assert_equal(sort_tuple_list([(2, 3), (1, 2), (3, 1)]), [(3, 1), (1, 2), (2, 3)])
assert_equal(sort_tuple_list([(1, 7), (1, 3), (3, 4, 5), (2, 2)]), [(2, 2), (1, 3), (3, 4, 5), (1, 7)])
```

# 5. (1.0 point) List sort

Given two lists sorted in increasing order, create and return a merged list of all the elements in sorted order. You may modify the passed in lists. Ideally, the solution should work in "linear" time, making a single pass of both lists.

<!-- grade_id: Q5A, points: 0, grade: false, solution: true -->
```python
def list_merge(list1, list2):
    '''Given two lists sorted in increasing order, create and 
       return a merged list of all the elements in sorted orde
    '''
    return sorted(list1 + list2)
    ### BEGIN SOLUTION
    ### END SOLUTION
```

<!-- grade_id: Q5B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(list_merge(['aa', 'xx', 'zz'], ['bb', 'cc']), ['aa', 'bb', 'cc', 'xx', 'zz'])
assert_equal(list_merge(['aa', 'xx'], ['bb', 'cc', 'zz']), ['aa', 'bb', 'cc', 'xx', 'zz'])
assert_equal(list_merge(['aa', 'aa'], ['aa', 'bb', 'bb']), ['aa', 'aa', 'aa', 'bb', 'bb'])
list_merge(['aa', 'ccc'], ['aa', 'bb', 'bb'])
```

Output:
```
['aa', 'aaa', 'bb', 'bb', 'ccc']
```

# 6. (1.0 point) Built-in Operator

Imagine that Python doesn’t have the `!=` operator built in. Write a function that takes two parameters and gives the same result as the `!=` operator. Obviously, you cannot use `!=` within your function! Test if your code works by thinking of examples and making sure the output is the same for your new method as `!=` gives you.

```python
rand_string = 'This String.'

print(sorted(set(rand_string)))
print(type(sorted(rand_string)), sorted(rand_string))
print(sorted(rand_string) - sorted(rand_string))
```

Output:
```
[' ', '.', 'S', 'T', 'g', 'h', 'i', 'n', 'r', 's', 't']
<class 'list'> [' ', '.', 'S', 'T', 'g', 'h', 'i', 'i', 'n', 'r', 's', 't']
TypeError                                 Traceback (most recent call last)
<ipython-input-53-da11d0cf9284> in <module>()
      3 print(sorted(set(rand_string)))
      4 print(type(sorted(rand_string)), sorted(rand_string))
----> 5 print(sorted(rand_string) - sorted(rand_string))

TypeError: unsupported operand type(s) for -: 'list' and 'list'
```

<!-- grade_id: Q6A, points: 0, grade: false, solution: true -->
```python
    ### BEGiN SOLUTION 1
    
def not_equal(val1, val2):
    ''' User supplied version of the != operator, implemeted without using !='''
    ### BEGIN SOLUTION
    
    # for numbers
    if val1 == val2:
        return False
    else:
        return True
    
    
    
### END SOLUTION 1


### BEGIN SOLUTION 2 with no != or ==

def not_equaltwo(val1, val2):
    ''' User supplied version of the != operator, implemeted without using !='''

    
    # for numbers
    if val1 > val2 or val1 < val2:
        return True
    
    # for strings
    elif sorted(val1) == sorted(val2):
        return False
    
    # for boolean vals 
    
    
    else:
        return False
    
    ### END SOLUTION 2
    
    

    
    
    
```

<!-- grade_id: Q6B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_true, assert_false
assert_true(not_equal(2,3))
assert_false(not_equal(3,3))
assert_true(not_equal('3',3))
assert_true(not_equal((1,2),[1,2]))
assert_false(not_equal((3,4),(3,4)))
```

## 7. (1.5 point) Quadratic Formula

A quadratic equation is a second order polynomial with three coefficients a, b, and c. It is given by:

$$ax^2 + bx + c = 0$$

The solutions to the equation are given by the formulas:

$$x_1 = \frac{-b + \sqrt{b^2 - 4ac}}{2a}$$

and

$$x_2 = \frac{-b - \sqrt{b^2 - 4ac}}{2a}$$



Write a function that computes the roots of a quadratic equation and returns them as a tuple. Your function should take three parameters, a, b, and c. Note that if

$$b^2 - 4ac < 0$$ 

then the roots will be complex. Hint: remember that Python has a built-in complex data type.

<!-- grade_id: Q7A, points: 0, grade: false, solution: true -->
```python
import cmath
import math
import numpy as np 

print(math.sqrt(9))
print(complex(22/7))

def roots(a_in, b_in, c_in):
    ''' Computes the roots of the quadratic equation with coefficients a, b, c, and returns
        them as the tuple (x1, x2)
    '''
    ### BEGIN SOLUTION  
    
    a = complex(a_in)
    b = complex(b_in)
    c = complex(c_in)
    
    x1 = ((-b + cmath.sqrt(b**2 - 4*a*c))/(2*a))
    x2 = ((-b - cmath.sqrt(b**2 - 4*a*c))/(2*a))
    
    return (x1, x2)
    
    
    
    ### END SOLUTION
```

Output:
```
3.0
(3.142857142857143+0j)
```

<!-- grade_id: Q7B, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(roots(1,2,3), ((-1+1.4142135623730951j), (-1-1.4142135623730951j)))
assert_equal(roots(0,0,4), (0j, 0j))
assert_equal(roots(-4,3,2), (-0.42539052967910607, 1.175390529679106))
```

Output:
```
<ipython-input-87-b1e16debcc47> in <module>()
      1 from nose.tools import assert_equal
      2 assert_equal(roots(1,2,3), ((-1+1.4142135623730951j), (-1-1.4142135623730951j)))
----> 3 assert_equal(roots(0,0,4), (0j, 0j))
      4 assert_equal(roots(-4,3,2), (-0.42539052967910607, 1.175390529679106))

<ipython-input-86-ad5e8b388b08> in roots(a_in, b_in, c_in)
     15     c = complex(c_in)
     16 
---> 17     x1 = ((-b + cmath.sqrt(b**2 - 4*a*c))/(2*a))
     18     x2 = ((-b - cmath.sqrt(b**2 - 4*a*c))/(2*a))
     19 

ZeroDivisionError: complex division by zero
```

## 8. (2.5 points) The Game of Nims

In this game, two players sit in front of a pile of 100 stones. They take turns, each removing between 1 and 5 stones (assuming there are at least 5 stones left in the pile). The person who removes the last stone(s) wins.

In this problem, you’ll write a function to play this game. It may seem tricky, so break it down into parts. Like many programs, we have to use nested loops (one loop inside another). In the outermost loop, we want to keep playing until we are out of stones. Inside that, we want to keep alternating players. You have the option of either writing two blocks of code, or keeping a variable that tracks the current player.

We are going to use the `input` function to obtain the user's move. We might want to have an innermost loop that checks if the user’s input is valid. Is it a number? Is it a valid number (e.g. between 1 and 5)? Are there enough stones in the pile to take off this many? If any of these answers are no, we should tell the user and re-ask them the question.

If you choose to write two blocks of code, the basic outline of the program should be something like this:

    while [pile is not empty]:
       while [player 1’s answer is not valid]:
          [ask player 1]
       [execute player 1’s move]
       [same as above for player 2]
       
Be careful with the validity checks. Specifically, we want to keep asking player 1 for their choice as long as their answer is not valid, BUT we want to make sure we ask them at least ONCE. So, for example, we will want to keep a variable that tracks whether their answer is valid, and set it to `False` initially.

There is no need to limit yourself to only one function! Use as many functions as you want in order to make the code more readable and modular.

When you're finished, test each other's programs by playing them!

<!-- grade_id: Q8A, points: 4, grade: true, solution: true -->
```python
def play_nims(pile, max_stones):
    '''
    An interactive two-person game; also known as Stones.
    @param pile: the number of stones in the pile to start
    @param max_stones: the maximum number of stones you can take in one turn
    '''
    #turn check into loop? 
    if str(type(max_stones)) in ["<class 'float'>","<class 'int'>"]:
        pass
    else:
        turn_size = input('Sorry, please enter a number for max turn size ->')
        
    if str(type(pile)) in ["<class 'float'>","<class 'int'>"]:
        pass
    else:
        total_pile = input('Sorry, please enter a number total stones ->')
    
    
## Basic structure of program (feel free to alter as you please):

#    while [pile is not empty]:
#        while [player 1's answer is not valid]:
#            [ask player 1]
#            [execute player 1's move]
#       
#        while [player 2's answer is not valid]:
#            [ask player 2]
#            [execute player 2's move]
#
#    print "Game over"

### BEGIN SOLUTION
    winning_player = 'unk'
    total_pile = pile
    turn_size = max_stones
    
    while total_pile > turn_size:
        
        #player1 turn

        p1_turn = input("There are %d stones remaining. \nPlayer 1 choose how many stones to remove from the pile -> " %total_pile)

        while True:
                    try:
                        int(p1_turn)
                        assert(int(p1_turn) <= turn_size)
                        break
                    except ValueError:
                        p1_turn = input("Sorry player 1, you must enter an interger -> ")
                    except AssertionError:
                        p1_turn = input("Sorry player 1, your turn must be less than the max turn size -> ")
        
        p1_turn = int(p1_turn)
        total_pile -= p1_turn
        winning_player = 'Player 2'
        
        if total_pile <= turn_size:
            break
        else:
            pass

        # player2 turn

        p2_turn = input("There are %d stones remaining. \nPlayer 2 choose how many stones to remove from the pile -> " %total_pile)

        while True:
                    try:
                        int(p2_turn)
                        assert(int(p2_turn) <= turn_size)
                        break
                    except ValueError:
                        p2_turn = input("Sorry player 2, you must enter an interger -> ")
                    except AssertionError:
                        p2_turn = input("Sorry player 2, your turn must be less than the max turn size -> ")
        
        p2_turn = int(p2_turn)
        total_pile -= p2_turn
        winning_player = 'Player 1'

    print('\n', winning_player, ' wins!')






### END SOLUTION
```

```python
play_nims(20,5)
```

Output:
```
/Users/Sam_I_Am/anaconda/lib/python3.6/site-packages/ipykernel/kernelbase.py in raw_input(self, prompt)
    692             self._parent_ident,
    693             self._parent_header,
--> 694             password=False,
    695         )
    696 

/Users/Sam_I_Am/anaconda/lib/python3.6/site-packages/ipykernel/kernelbase.py in _input_request(self, prompt, ident, parent, password)
    722             except KeyboardInterrupt:
    723                 # re-raise KeyboardInterrupt, to truncate traceback
--> 724                 raise KeyboardInterrupt
    725             else:
    726                 break

KeyboardInterrupt:
```

```python
play_nims(20,5)
```

Output:
```
There are 20 stones remaining. 
Player 1 choose how many stones to remove from the pile -> 5
There are 15 stones remaining. 
Player 2 choose how many stones to remove from the pile -> 4
There are 11 stones remaining. 
Player 1 choose how many stones to remove from the pile -> 5
There are 6 stones remaining. 
Player 2 choose how many stones to remove from the pile -> 1

 Player 1  wins!
```

```python

```